Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Enjui Chang"
COLLABORATORS = "Li-lian Ang"
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    n = len(A) #find the length of the list
    maxsum = -float("inf") #the default maxsum is -∞
    
    #for loop for iterating 1 to the length of the lsit
    for i in range(n): 
        temp=0 #reset the temporary sum to 0
        
        #for loop that sums from i to the last element
        for j in range(i,n): 
            temp+=A[j] #sum the current sum with the next number
            
            #if the temporary sum is larger than the maxsum,
            #register the index and the sum
            if temp>maxsum: 
                maxsum=temp 
                start = i
                end = j
                
    return(start,end,maxsum)
```

```python
A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
No, we don't need to rerun the entire array again. Since the maximum subarray would only be differnt if it includes the last element, we only need to start with the last element until it reaches 4 at index 2.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. It needs to include 8 since without 8, array B is exactly the same as array A, which means that the maximum subarray is 4. Therefore, having a different maximum subarray without 8 is contradictory.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        last element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    #______return the value for a list with the length of 1 element_____#
    if len(x) == 1:
        return x[0] #return the value of the sole element
    
    #______initialize the variables______#
    temp_sum = 0 #the temporary sum after each iteration
    temp_mx = -float("inf") #the temporary max after each iteration
    
    #______iterate the reversely to see if the maximum subarray value changes______#
    for i in range(len(x)-1,0,-1):
        temp_sum += x[i]
        
        #update the temp_max if it is bigger than the than the temp_sum
        if temp_sum > temp_mx:
            temp_mx = temp_sum
            
    return max(mx,temp_mx)#check if the new maximum is bigger than the original sum
    
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    #_______initialize the maximum subarray______#
    mx = A[0]
    
    #_______iterate through every list starting from the list containing only one 
    #element to the full list_________#
    for i in range(len(A)):
        
        #find the maximum subarray using incremental_max_subarray
        temp_mx = incremental_max_subarray(A[:i+1],mx) 
        
        #update the maximum subarray
        if temp_mx > mx:
            mx = temp_mx
    
    #return the final maximum subarray
    return mx
    
```

```python
A = [-2, -3, 4, -1, -2, 1, 8]
B = [0]
C = [-2, -3, 4, -1, -2, 1, 8, 10, -4, 10]

assert(max_subarray(A) == 10)
assert(max_subarray(B) == 0)
assert(max_subarray(C) == 26)
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
This is not more efficient than the divide-and-conquer approach since the time complexity of divide-and-conquer is $O(n\lg n)$, while the time complexity of the incremental maximum subarray is $O(n^2)$. The big-O notation for the incremental is induced by finding two for-loops in the fucntion. 

```python
def recursive_maxsubarray(x, mx = -float("inf")):
    n = len(x)
    if(n==1):
        if(x[0]>mx):
            return x[0]
       	else:
            return mx
    maxarr = mx
    for i in range(n):
        sumarr = 0
        for j in range(i,n):
            sumarr+=x[j]
        if(sumarr>maxarr):
            maxarr = sumarr
    return recursive_maxsubarray(x[0:-1],maxarr)

recursive_maxsubarray([1,2,3,4,5,-6])
```